Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE"

Do not edit or insert code into the test cells as all you insert will be overwritten by the automated testing code.

---

### Fit 3D lines to cylinders
This project aims at fitting lines in 3D to data consisting of cylinders (produced by individual Geiger counters). Each cylinder is hence centred around an individual wire (parallel to the z-axis) on a regular grid and has a fixed radius. The cylinder height corresponds to a fixed z-coordinate error. The radius also has an error (at 10% of the radius value or at least 0.5 mm).

You have access to a file containing 10 events. The format is defined below. Each event contains data corresponding to two
lines. Each line originates from a simulated electron, each starting at the `x=0` plane from the same point, one emitted to the
left, i.e. negative x-coordinates, one to the right, with positive x-coordinates. Coordinates are given in millimetres.

The file arrives in compressed form. You can read it as given using `np.loadtxt(filename, delimiter=',')`. The column entries (and type) are
- Event number (int, starting at value 1)
- Line number (int, starting at value 1)
- Wire x-coordinate (float)
- Wire y-coordinate (float)
- Wire z-coordinate (float)
- Radius (float)
- Radius error (float)
- Wire z-coordinate error (float)

An example event is shown in the figure below. For your own figure, note that the matplotlib.pyplot library has a `pyplot.Circle((x, y), radius, fill=False, lw=2)` function that can be added as a 'patch' to a figure.

![evt](SNLRtracks.png)

The formula describing the distance of a line to a cylinder is 

$$D = \left| (\vec{x}_{p}-\vec{x}_{0}) \times{}  \vec{u} \right| - r$$

where $\vec{u}$ is the unit direction vector of the line and $\vec{x}_{0}$ is a point on the line. This assumes the line is not parallel to the y-z plane like in this exercise. The residuals between model and data also come with a weighting factor, made from the data errors, $\delta{}r$ and $\delta{}z$ according to 

$$W = \left[\sqrt{ (\delta{}z\ u_{z})^{2} + (\delta{}z\ u_{y})^{2} + (\delta{}r\ u_{z})^{2} + (\delta{}z\ u_{x})^{2} + (\delta{}r\ u_{y})^{2} + (\delta{}r\ u_{x})^{2}}\right]^{-1}$$

where $u_{x,y,z}$ represents the three components of the unit vector.

**Project**: Write an application that 
1. reads the data and 
2. fits lines to all events, 
3. counts and returns the number of successful fits, 
4. stores your results in a file with the format given below and finally 
5. plots the first event following the plot format from the example picture above.

The given code in the cell should be incorporated into your solution since its function names will be used in the tests. The `runFit(infile, outfile)` function should steer the execution of the entire code with the benefit that it can be used to measure the execution time (will be done in the test, but also useful for yourself). The `readData(filename)` function is independent of the file content and can read the compressed files, for testing or later for assessment. Two data files will be supplied for the project for your own testing. One contains 10 events in the right data format and one contains truth information (as human readable file) for checking your own code. The model solution fits all 10 events correctly.

**Hint**: The data in the file has been created with an intercept in the z-coordinate at 0.0 in order to stabilize the fits a little and simplify your task. Note that you should not take that insight as a fixed parameter (would lead to a reduction of marks) but as a reassurance to test your attempts relatively easily. The assessment data file also has the creation intercept set at `z=0` for all events.

**Hint 2**: The errors on the intercept in z and the slope in the x-z plane are likely to be big, occasionally. Not to worry, this reflects the large error on the z-coordinate which can induce a broad minimum and hence big errors. This observation simply is in the nature of the data.

Output format for fit results:
- Event number (int, starting at 1)
- Line number (int, starting at value 1)
- intercept in y (float)
- error on intercept in y (float)
- slope in x-y plane (float)
- error on slope in x-y plane (float)
- intercept in z (float)
- error on intercept in z (float)
- slope in x-z plane (float)
- error on slope in x-z plane (float)

On storing results: it is important that every line in every event has a result, even for failed hits in order to enable the comparison for your assessment. A failed fit should therefore store a result as

```python
datastore.append((evtnr, linenr, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0))
```

with `evtnr` the event number (int) and `linenr` the line number (int).

#### Marking
In all likelihood, not all fits will be successful. The model solution achieves a 98% efficiency. Your code will be run on a larger data file with 1000 events on the assessment workstation. 10 out of 20 marks will be awarded for solutions reaching 98% or higher efficiency (with the caveat of manually checking the code for correctness, of course). Note that your code efficiency will be cross checked with the truth data, i.e. the slopes and intercepts which produced each event. A successful fit for the assessment will hence have to satisfy a comparison between your parameter fit solutions and the truth values within the errors of your fit results.

Further 5 marks are awarded for the correct plot, i.e. similar to the example picture, plot the x-y plane with data Circles and your fit solution line (in that 2D plane). The final marks are awarded for any successful effort on speeding up the execution time compared to the pure, serial Python solution. Fitting 1000 events takes a fairly long time (a bit more than 15 minutes on my humble machine) hence any successful speed optimization is appreciated with up to 5 marks. Optimization can bring the time down to 40 s but at the expense of efficiency. The 98% efficiency was achieved in about 95 s.

<!-- grade_id: FitCode, points: 18, grade: true, solution: true -->
```python
'''
Fit 3D lines to cylinders
'''

# YOUR CODE HERE

def readData(filename):
    '''
    Read compressed data.
    '''
    return np.loadtxt(filename, delimiter=",")


def runFit(inputfilename, outputfilename):
    '''

    Parameters
    ----------
    inputfilename : string
        input data file name for fitting.
    outputfilename : string
        result storage file name.

    Returns
    -------
    counter : int
        number of successful fits; 100% would be twice the number
        of events (two lines per event).

    '''
    fulldata = readData(inputfilename)

    counter = 0
    
# YOUR CODE HERE


    return counter

# run the code and plot the figure

# YOUR CODE HERE


```

<!-- grade_id: FittestA, points: 1, grade: true, solution: false -->
```python
# test cell can't work since reading a file is called that you don't have. Consider it example code.
# main script
import time
start = time.process_time()
success = runFit("ringdata.txt.gz", "output.txt.gz")
stop = time.process_time()
print ("success: ", success)
print ("fit time: ", (stop-start))

# testing part
truthdata  = readData("ringtruth.txt")
outputdata = readData("output.txt.gz")

# test 1
nevents = int(outputdata[-1,0])
assert nevents == int(len(truthdata) / 2) ,'equality of data volume test failed.'
```

<!-- grade_id: FittestB, points: 1, grade: true, solution: false -->
```python
# test 2: compare your fit results with the truth data and determine the fit efficiency.
# Again, can't run since you don't have the 1000 event data file for the assignment but
# may clarify what is meant by fit efficiency.
truthrow = 0
counter = 0
for evt in range(1,nevents+1):
    # select event number data
    evtdata = outputdata[outputdata[:,0]==evt]
    for i in range(2):
        # now check on values
        trueixy  = truthdata[truthrow, 0]
        trueslxy = truthdata[truthrow, 2]
        ixy  = evtdata[i, 2]
        slxy = evtdata[i, 4]
        errixy  = evtdata[i, 3]
        errslxy = evtdata[i, 5]
        # catch fit failed case
        if errixy==0.0 and errslxy==0.0:
            truthrow += 1 # next line
            continue # skip the rest below
        bool1 = trueixy <= ixy + errixy and trueixy >= ixy - errixy
        bool2 = trueslxy <= slxy + errslxy and trueslxy >= slxy - errslxy
        if bool1 and bool2:
            counter += 1
        else:
            print(bool1, bool2)
            print('True: ', trueixy, trueslxy)
            print('Fit 1: ', ixy, slxy)
            print('Error 1: ', errixy, errslxy)
        truthrow += 1 # next line
print('efficiency: ', counter/(2*nevents))
assert counter/(2*nevents)>=0.98, 'Less than 98% efficiency. Reduce marks'
```